## Problem 3 - Applying the classifier (*4 points*)

Finally, we can bring together the pieces that we have created in problems 1 and 2. Here, your aim is to take
advantage of your new functions and sort a dataset of temperatures in Fahrenheit into four different classes.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable, function, and parameter names as stated in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions. Also, please do not copy/paste existing code cells. You can add cells if you like, but copying and pasting can also cause the automatic tests to fail.

**Your score on this problem will be based on following criteria:**

- Creating a new **script** file called `temp_functions.py` with the functions that you wrote in Problems 1 and 2
- Add a docstring to the script file
- Converting and classifying the list of temperatures in Part 2
- Storing the number of temperatures in each class in the specified variables
- Answering some questions at the end of this problem
- Uploading your notebook to your GitHub repository for this week's exercise

### Part 1: Creating a script file (1 point)

For this part you should

- Create a new **script** file called `temp_functions.py` in your `exercise-4-username` directory as we saw in the Lesson 4 materials. Copy and paste the functions that you wrote in Problems 1 and 2 into that script file (i.e., the `fahr_to_celsius` and `temp_classifier` functions)
- Add a docstring to the script file (in addition to the separate docstrings in the functions)
- Save the script file (CTRL + S) and continue working in **this** Notebook (i.e., `Exercise-4-problem-3.ipynb`)

### Part 2: Convert Fahrenheit temperatures to Celsius (3 points)

For this part you should

1. Store the Fahrenheit temperatures below into the list `temp_data` (just execute the existing code cell below).

```python
# List of half-hourly temperature values (in degrees Fahrenheit) for one week
temp_data =  [19, 21, 21, 21, 23, 23, 23, 21, 19, 21, 19, 21, 23, 27, 27, 28, 30, 30, 32, 32, 32, 32, 
              34, 34, 34, 36, 36, 36, 36, 36, 36, 34, 34, 34, 34, 34, 34, 32, 30, 30, 30, 28, 28, 27,
              27, 27, 23, 23, 21, 21, 21, 19, 19, 19, 18, 18, 21, 27, 28, 30, 32, 34, 36, 37, 37, 37, 
              39, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 39, 39, 37, 37, 36, 36, 34, 34, 32, 30, 30,
              28, 27, 27, 25, 23, 23, 21, 21, 19, 19, 19, 18, 18, 18, 21, 25, 27, 28, 34, 34, 41, 37, 
              37, 39, 39, 39, 39, 41, 41, 39, 39, 39, 39, 39, 41, 39, 39, 39, 37, 36, 34, 32, 28, 28,
              27, 25, 25, 25, 23, 23, 23, 23, 21, 21, 21, 21, 19, 21, 19, 21, 21, 19, 21, 27, 28, 32,
              36, 36, 37, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 41, 41, 41, 41, 39, 37, 36, 36, 34,
              32, 30, 28, 28, 27, 27, 25, 25, 23, 23, 23, 21, 21, 21, 19, 19, 19, 19, 19, 19, 21, 23,
              23, 23, 25, 27, 30, 36, 37, 37, 39, 39, 41, 41, 41, 39, 39, 41, 43, 43, 43, 43, 43, 43,
              43, 43, 43, 39, 37, 37, 37, 36, 36, 36, 36, 34, 32, 32, 32, 32, 30, 30, 28, 28, 28, 27,
              27, 27, 27, 25, 27, 27, 27, 28, 28, 28, 30, 32, 32, 32, 34, 34, 36, 36, 36, 37, 37, 37,
              37, 37, 37, 37, 37, 37, 36, 34, 30, 30, 27, 27, 25, 25, 23, 21, 21, 21, 21, 19, 19, 19,
              19, 19, 18, 18, 18, 18, 18, 19, 23, 27, 30, 32, 32, 32, 32, 32, 32, 34, 34, 34, 34, 34,
              36, 36, 36, 36, 36, 32, 32, 32, 32, 32, 32, 32, 32, 30, 30, 30, 30, 30, 30, 30, 30, 30,
              30, 30, 30, 30, 28, 28]
```

#### Part 2 (continues)

2. Import the `fahr_to_celsius` and `temp_classifier` functions from your `temp_functions.py` 
3. Create an empty list called `temp_classes` (which will be filled with temperature class numbers later)
4. **Convert the Fahrenheit temperatures** in the `temp_data` list into Celsius and **classify the Celsius temperatures** using the importend functions. The output should be stored in the `temp_classes` list
 
   **You can do the temperature conversion and classification inside one for loop:**

    - Iterate over the list of Fahrenheit temperatures (create a `for` loop), and inside the loop:

        1. Create a new variable called `temp_celsius` in which you should assign the temperature in Celsius using the `fahr_to_celsius` function to convert the Fahrenheit temperature into Celsius.
        2. Create a new variable called `temp_class` in which you should assign the temperature class number (0, 1, 2, or 3) using the `temp_classifier` function
        3. Add the `temp_class` value to the `temp_classes` list

<!-- grade_id: problem_3_ag_p2_functions_usage, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
from temp_functions import fahr_to_celsius
from temp_functions import temp_classifier
temp_classes = []
type(temp_classes)
n = len(temp_data)

# LOOP
for temp in temp_data:
    temp_celsius = fahr_to_celsius(temp)
    temp_classes.append(temp_celsius)
    temp_class = temp_classifier(temp_celsius)
    temp_classes.append(temp_class)

print(temp_classes)
```

Output:
```
[-7.222222222222222, 0, -6.111111111111111, 0, -6.111111111111111, 0, -6.111111111111111, 0, -5.0, 0, -5.0, 0, -5.0, 0, -6.111111111111111, 0, -7.222222222222222, 0, -6.111111111111111, 0, -7.222222222222222, 0, -6.111111111111111, 0, -5.0, 0, -2.7777777777777777, 0, -2.7777777777777777, 0, -2.2222222222222223, 0, -1.1111111111111112, 1, -1.1111111111111112, 1, 0.0, 1, 0.0, 1, 0.0, 1, 0.0, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 2.2222222222222223, 2, 2.2222222222222223, 2, 2.2222222222222223, 2, 2.2222222222222223, 2, 2.2222222222222223, 2, 2.2222222222222223, 2, 1.1111111111111112, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 1.1111111111111112, 1, 0.0, 1, -1.1111111111111112, 1, -1.1111111111111112, 1, -1.1111111111111112, 1, -2.2222222222222223, 0, -2.2222222222222223, 0, -2.7777777777777777, 0, -2.7777777777777777, 0, -2.7777777777777777, 0, -5.0, 0, -5.0, 0, -6.111111111111111, 0, -6.111111111111111
```

```python
print(temp_class)

```

Output:
```
0
```

#### Part 2 (continues)

6. Calculate how many temperatures there are in each temperature class:

    1. Create four variables called `zeros`, `ones`, `twos`, and `threes` 
    2. Count and assign to each variable how many times values 0, 1, 2, and 3 are present in the `temp_classes` list and print out the results below. In other words, assign the value of how many time `0` is in the list to the variable `zeros`, and so on. 
  
**TIP**: You might want to consider using a [**count()** function](https://www.tutorialspoint.com/python3/list_count.htm) OR a for loop for this.

```python
#LOOP
for i in temp_classes:
    zeros = temp_classes.count(temp_classes == 0)
    ones = temp_classes.count(temp_classes == 1)
    twos = temp_classes.count(temp_classes == 2)
    threes = temp_classes.count(temp_classes == 3)



```

Output:
```
672
```

```python
# 1. How many 0 values exist in the temp_classes list?
print(zeros)
# 2. How many 1 values exist in the temp_classes list?
print(ones)
# 3. How many 2 values exist in the temp_classes list?
print(twos)
# 4. How many 3 values exist in the temp_classes list?
print(threes)
```

Output:
```
168
168
168
168
```

### Problem 3 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_3_ag_p2_functions_usage_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that functions are in the namespace
assert inspect.isfunction(fahr_to_celsius)
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that variable has been created
assert 'temp_celsius' in locals()
assert 'temp_class' in locals()
assert 'temp_classes' in locals()

# Check that temp_classes is a list
assert type(temp_classes) == list
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test3, points: 0.5, grade: true, solution: false -->
```python
# Check that the functions have a single parameter
t_params = list(inspect.signature(temp_classifier).parameters.keys())
f_params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(t_params) == 1
assert len(f_params) == 1
```

Check that required variables exists and print their value (check manually that the answers make sense!):

<!-- grade_id: problem_3_ag_p2_functions_usage_test4, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "zeros" 
assert 'zeros' in locals()
print(zeros)
```

Output:
```
168
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test5, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "ones" 
assert 'ones' in locals()
print(ones)
```

Output:
```
168
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test6, points: 0.5, grade: true, solution: false -->
```python
# Checkthe variable "twos" 
assert 'twos' in locals()
print(twos)
```

Output:
```
168
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test7, points: 0.5, grade: true, solution: false -->
```python
# Check the variable "threes" 
assert 'threes' in locals()
print(threes)
```

Output:
```
168
```

### Check your code
 
- Remeber to check that your have commented your code and used the required variable names. 
- Make sure to commit and push also your `temp_functions.py` script to your own GitHub repository for Exercise 4!
- If you have any comments and concerns at this point, you can write them below:

<!-- grade_id: problem_3_mg_p3_check_your_code, points: 0.5, grade: true, solution: true -->
YOUR ANSWER HERE

#### All done!

That's it! Now you have finished all the problems for Exercise 4! Remember to push all your changes to your personal GitHub Exercise 4 repository (including the `temp_functions.py` script).