## Problem 2 - Temperature classifier (*3 points*)

In [Problem 2 of Exercise 3 last week](https://github.com/Geo-Python-2020/Exercise-3/blob/master/Exercise-3-problem-2.ipynb) we practiced how to classify temperatures into four different classes (*i.e., cold, slippery, comfortable, and warm*). Let's continue working with the same idea. 

Here, you should create a function called `temp_classifier` that accepts a temperature value in Celsius that will be reclassified into integer numbers 0-3 based on following criteria:

| Return value | Classification criteria                                                  |
| :----------: | :----------------------------------------------------------------------- |
| 0            | Temperatures below -2 degrees Celsius                                    |
| 1            | Temperatures equal or warmer than -2, but less than +2 degrees Celsius   |
| 2            | Temperatures equal or warmer than +2, but less than +15 degrees Celsius  |
| 3            | Temperatures equal or warmer than +15 degrees Celsius                    |

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable, function, and parameter names as stated in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions. Also, please do not copy/paste existing code cells. You can add cells if you like, but copying and pasting can also cause the automatic tests to fail.

**Your score on this problem will be based on following criteria:**

- Creating a new function called `temp_classifier` **(0.5 points)**
- Defining one parameter called `temp_celsius` in your `temp_classifier` function **(0.3 points)**
- Reclassifying input temperatures based on the criteria in the above table and returning the reclassified value as an integer number (i.e., 0, 1, 2, or 3) **(1.2 points)**
- Adding comments in your code and a docstring that explains how to use your `temp_classifier` function (i.e., you should write the purpose of the function, parameters, and returned values) **(1.0 points)**
- Uploading your notebook to your GitHub repository for this week's exercise

<!-- grade_id: problem_2_ag_p1_temp_classifier, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def temp_classifier(temp_celsius):
    if temp_celsius < -2:
       value = 0
    elif 2 > temp_celsius >= - 2:
        value = 1
    elif 15 > temp_celsius >= 2:
        value = 2
    elif temp_celsius >= 15:
        value = 3
    return(value)

    




```

**Use** your newly created function to print the answers to the following questions:

```python
# 1. What is the class value for 16.5 degrees Celsius?
print(temp_classifier(16.5))

# 2. What is the class value for +2 degrees Celsius?
print(temp_classifier(2))

```

Output:
```
3
2
```

### Problem 2 tests

Run the code cells below to check your answer. Note that you are not able to modify the test cells. Your solution is on the right track if you don't see any errors when running these code cells.

<!-- grade_id: problem_2_ag_p1_temp_classifier_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Check that function exists
assert inspect.isfunction(temp_classifier)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test2, points: 0.3, grade: true, solution: false -->
```python
# Check that the function has a single parameter and the pamameter name is correct
params = list(inspect.signature(temp_classifier).parameters.keys())
assert len(params) == 1
assert params[0] == 'temp_celsius'
```

Check that the function produces correct answers for selected values:

<!-- grade_id: problem_2_ag_p1_temp_classifier_test3, points: 0.3, grade: true, solution: false -->
```python
# 1. What is the class value for 16.5 degrees (Celsius)?
assert temp_classifier(16.5) == 3, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test4, points: 0.3, grade: true, solution: false -->
```python
# 2. What is the class value for +2 degrees (Celsius)?
assert temp_classifier(2) == 2, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test5, points: 0.3, grade: true, solution: false -->
```python
# 3. What is the class value for +1 degrees (Celsius)?
assert temp_classifier(1) == 1, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

<!-- grade_id: problem_2_ag_p1_temp_classifier_test6, points: 0.3, grade: true, solution: false -->
```python
# 4. What is the class value for -5 degrees (Celsius)?
assert temp_classifier(-5) == 0, 'Wrong class'
print("ok :)")
```

Output:
```
ok :)
```

### Check your code

 - Are all the variable names are correct?
 - Does your code has enough comments?
 - Does your function has an informative dosctring?

<!-- grade_id: problem_2_manual_check, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

#### Done!

That's it! Now you are ready to continue with [Problem 3](Exercise-4-problem-3.ipynb).